# Functions

References:

[1] B Downey, A. (2012). *Think Python: How to Think Like a Computer Scientist-2e.*

[2] Gries, P., Campbell, J., & Montojo, J. (2017). *Practical programming: an introduction to computer science using Python 3.6.* Pragmatic Bookshelf.

[3] Matthes, E. (2023). *Python crash course: A hands-on, project-based introduction to programming.*

## 1 Functions in Python

One way to improve the modularity of your code is to adhere to the **DRY** principle, which means *Don't repeat yourself*. In other words, whenever you found yourself doing the same thing over and over again, do consider replacing it with abstraction or modules to avoid redundancy in your code.

### Defining functions in Python

The general form of a function definition is as follows:

```python
def <<function_name>>(<<arguments>>):
    <<function_body>>
```

The first line of the function definition is called the *function header*. which contains the *function name* and its corresponding *arguments*. Arguments are expressions that appears between the parenthesis of a function that the function needs to execute the code block defined inside it. We input it upon *calling* the function.

In our first notebook, we created a function that converts a given Celsius temperature to Fahrenheit temperature. If we are in a situation where we would always need to convert Celsius temperature to Fahrenheit temperature, say we are creating a fitness app and would like to accomodate both units of temperature, we can create a function for it.

```python
def convert_to_fahrenheit(celsius):
    """Convert a given temperature in celsius to fahrenheit"""
    fahrenheit = 9/5*celsius + 32

    return fahrenheit

def convert_to_celsius(fahrenheit):
    """Convert a given temperature in fahrenheit to celsius"""
    celsius = (fahrenheit - 32) * 5/9

    return celsius
```

Defining a function creates a function object, which has the type `function`

```python
type(convert_to_celsius)
```

Output:
```
function
```

You then call the function by using its name together with the required parameters enclosed in a parenthesis

```python
convert_to_celsius(98.6)
```

Output:
```
37.0
```

```python
convert_to_fahrenheit(29.0)
```

Output:
```
84.2
```

Most function definitions have a `return` statement, that when executed, ends the function and produces a *value*.

```python
return <<expression>>
```

When Python executes a `return` statement, it evaluates the expression then produces the result of that expression as the result of the function call.

### Passing arguments

When you pass arguments to a function, Python matches those arguments to the function's parameters depending on the way you specify them. You can pass arguments in a number of ways. You can use *positional arguments*, which need to be in the same order in which the parameters were written; *keyword arguments*, where each argument constists of a variable name and value.

```python
def describe_pet(animal_type, pet_name):
    """Display information about a pet."""
    print(f"I have a {animal_type}.")
    print(f"My {animal_type}'s name is {pet_name.title()}.")
```

```python
describe_pet('dog', 'Maki')
```

Output:
```
I have a dog.
My dog's name is Maki.
```

```python
describe_pet(animal_type='cat', pet_name='Midnight')
```

Output:
```
I have a cat.
My cat's name is Midnight.
```

```python
describe_pet('Maki', 'dog')
```

Output:
```
I have a Maki.
My Maki's name is Dog.
```

```python
describe_pet(pet_name='Midnight', animal_type='cat')
```

Output:
```
I have a cat.
My cat's name is Midnight.
```

Sometimes it makes sense to give an argument a default value, for example, in a scenario where you want to make the argument optional.

```python
def get_formatted_name(first_name, last_name, middle_name=''):
    """Return full name, neatly formatted written last name first."""
    if middle_name:
        full_name = f"{last_name}, {first_name} {middle_name}"
    else:
        full_name = f"{last_name}, {first_name}"

    return full_name.title()
```

```python
get_formatted_name("Leodegario", "Lorenzo")
```

Output:
```
'Lorenzo, Leodegario'
```

```python
get_formatted_name("Leodegario", "Lorenzo", "Verso")
```

Output:
```
'Lorenzo, Leodegario Verso'
```

## 2 Designing New Functions: A Recipe

Everytime you write a function, do figure out the answers to the following questions:

1. What is an appropriate and informative name for the function?
2. What are the parameters, and what types of information do they refer to?
3. What calculations are you doing with that information?
4. What information does the function return?
5. Does it work like you expect it to?

Part of the outcome upon completing this questions should be a working function. But equally important is the *documentation* for the function. Use three double quotes `"""` to start and end this documentation. This notation is called the *docstring* or the *documentation string*.

Let's improve our temperature converter functions according to this framework.

*Note: the docstring format we specifically use is the [numpy](https://numpydoc.readthedocs.io/en/latest/format.html) docstring format. Which is the common notation for many data science libraries that you'll eventually be using*

```python
def convert_to_fahrenheit(celsius: float) -> float:
    """Convert a given temperature in celsius to fahrenheit

    Parameters
    ----------
    celsius : float
        Input temperature in celsius

    Returns
    -------
    fahrenheit : float
        Temperature of the input in fahrenheit

    Examples
    --------
    >>> convert_to_fahrenheit(37.0)
    98.6
    """
    fahrenheit = 9/5*celsius + 32

    return fahrenheit

def convert_to_celsius(fahrenheit: float) -> float:
    """Convert a given temperature in fahrenheit to celsius

    Parameters
    ----------
    fahrenheit : float
        Input temperature in fahrenheit

    Returns
    -------
    celsius : float
        Temperature of the input in celsius

    Examples
    --------
    >>> convert_to_celsius(98.6)
    37.0
    """
    celsius = (fahrenheit - 32) * 5/9

    return celsius
```

### Functions That Python Provides

To improve also the readability and conciseness of your code, consider using built-in functions to perform common operations. We've already encountered some of these functions in the previous session which includes `print`, `float`, `int`, `help`, and `input`.

Found [here](https://www.w3schools.com/python/python_ref_functions.asp) are other common built-in functions in Python: 

```python
abs(-9)
```

Output:
```
9
```

```python
round(1.337, 2)
```

Output:
```
1.34
```

```python
help(round)
```

Output:
```
Help on built-in function round in module builtins:

round(number, ndigits=None)
    Round a number to a given precision in decimal digits.

    The return value is an integer if ndigits is omitted or None.  Otherwise
    the return value has the same type as the number.  ndigits may be negative.
```

```python
pow(25, 0.5)
```

Output:
```
5.0
```

## 3 A Modular Approach to Program Organization

> *Mathematicians don't prove every theorem from scratch. Instead, they build their proofs on truths their predecessors have already established. In the same way, it's rare for someone to write all of a program alone; it's much more common-and productive to make use of the millions of lines of code that other programmers have written before*

In Python, a *module* is a collection of variables and functions grouped together in a single file. These variables and functions usually help you accomplish a certain task. For example, the `math` module contains variables such as `pi` and mathematical functions such as `cos` (cosine) and `sqrt` (square root).

### Importing Modules

To gain access to the variables and functions defined in a module, we use the `import` statement.

#### `math` module

```python
import math
```

```python
type(math)
```

Output:
```
module
```

```python
help(math)
```

Output:
```
Help on module math:

NAME
    math

MODULE REFERENCE
    https://docs.python.org/3.12/library/math.html

    The following documentation is automatically generated from the Python
    source files.  It may be incomplete, incorrect or include features that
    are considered implementation detail and may vary between Python
    implementations.  When in doubt, consult the module reference at the
    location listed above.

DESCRIPTION
    This module provides access to the mathematical functions
    defined by the C standard.

FUNCTIONS
    acos(x, /)
        Return the arc cosine (measured in radians) of x.

        The result is between 0 and pi.

    acosh(x, /)
        Return the inverse hyperbolic cosine of x.

    asin(x, /)
        Return the arc sine (measured in radians) of x.

        The result is between -pi/2 and pi/2.

    asinh(x, /)
        Return the inverse hyperbolic sine of x.

    atan(x, /)
        Return the arc tangent (measured in radians) of x.

        The re
```

```python
math.sqrt(9)
```

Output:
```
3.0
```

## 4 Hands-on Exercises

### City names

Write a function called `city_country()` that takes in a name of a `city` and its `country`. The function should return a string formatted like this:

```
"Makati City, Philippines"
```

<!-- grade_id: cell-667d20c714a10725, points: 0, grade: false, solution: true -->
```python
def city_country(city, country):
    # YOUR CODE HERE
```

<!-- grade_id: cell-4359f1074a7acadd, points: 0, grade: true, solution: false -->
```python
assert isinstance(city_country('Makati City', 'Philippines'), str)
```

### Sum digits

Given a number `n`, create a function that computes the sum of the digits of this number.

<!-- grade_id: cell-69f87e694fdcb34e, points: 0, grade: false, solution: true -->
```python
def sum_digits(n):
    # YOUR CODE HERE
```

<!-- grade_id: cell-0f92ff5b7357b1a9, points: 0, grade: true, solution: false -->
```python
assert isinstance(sum_digits(1345), int)
```